# Module 09: Spatial - Processing Images with Convolutions

Welcome to Module 09! You'll implement spatial operations that transform machine learning from working with simple vectors to understanding images and spatial patterns.

## ðŸ”— Prerequisites & Progress
**You've Built**: Complete training pipeline with MLPs, optimizers, and data loaders
**You'll Build**: Spatial operations - Conv2d, MaxPool2d, AvgPool2d for image processing
**You'll Enable**: Convolutional Neural Networks (CNNs) for computer vision

**Connection Map**:
```
Training Pipeline â†’ Spatial Operations â†’ CNN (Milestone 03)
    (MLPs)            (Conv/Pool)        (Computer Vision)
```

## Learning Objectives
By the end of this module, you will:
1. Implement Conv2d with explicit loops to understand O(NÂ²MÂ²KÂ²) complexity
2. Build pooling operations (Max and Average) for spatial reduction
3. Understand receptive fields and spatial feature extraction
4. Analyze memory vs computation trade-offs in spatial operations

Let's get started!

## ðŸ“¦ Where This Code Lives in the Final Package

**Learning Side:** You work in `modules/09_spatial/spatial_dev.py`  
**Building Side:** Code exports to `tinytorch.core.spatial`

```python
# How to use this module:
from tinytorch.core.spatial import Conv2d, MaxPool2d, AvgPool2d
```

**Why this matters:**
- **Learning:** Complete spatial processing system in one focused module for deep understanding
- **Production:** Proper organization like PyTorch's torch.nn.Conv2d with all spatial operations together
- **Consistency:** All convolution and pooling operations in core.spatial
- **Integration:** Works seamlessly with existing layers for complete CNN architectures

<!-- grade_id: spatial-setup, points: 0, grade: false, solution: true -->
```python


#| default_exp core.spatial

#| export
import numpy as np

from tinytorch.core.tensor import Tensor
```

## 1. Introduction - What are Spatial Operations?

Spatial operations transform machine learning from working with simple vectors to understanding images and spatial patterns. When you look at a photo, your brain naturally processes spatial relationships - edges, textures, objects. Spatial operations give neural networks this same capability.

### The Two Core Spatial Operations

**Convolution**: Detects local patterns by sliding filters across the input
**Pooling**: Reduces spatial dimensions while preserving important features

### Visual Example: How Convolution Works

```
Input Image (5Ã—5):        Kernel (3Ã—3):        Output (3Ã—3):
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”      â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”         â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ 1 2 3 4 5 â”‚      â”‚ 1 0 -1 â”‚         â”‚ ? ? ? â”‚
â”‚ 6 7 8 9 0 â”‚  *   â”‚ 1 0 -1 â”‚    =    â”‚ ? ? ? â”‚
â”‚ 1 2 3 4 5 â”‚      â”‚ 1 0 -1 â”‚         â”‚ ? ? ? â”‚
â”‚ 6 7 8 9 0 â”‚      â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜         â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
â”‚ 1 2 3 4 5 â”‚
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜

Sliding Window Process:
Position (0,0): [1,2,3]   Position (0,1): [2,3,4]   Position (0,2): [3,4,5]
               [6,7,8] *               [7,8,9] *               [8,9,0] *
               [1,2,3]                 [2,3,4]                 [3,4,5]
               = Output[0,0]           = Output[0,1]           = Output[0,2]
```

Each output pixel summarizes a local neighborhood, allowing the network to detect patterns like edges, corners, and textures.

### Why Spatial Operations Transform ML

```
Without Convolution:                    With Convolution:
32Ã—32Ã—3 image = 3,072 inputs          32Ã—32Ã—3 â†’ Conv â†’ 32Ã—32Ã—16
â†“                                      â†“                     â†“
Dense(3072 â†’ 1000) = 3M parameters    Shared 3Ã—3 kernel = 432 parameters
â†“                                      â†“                     â†“
Memory explosion + no spatial awareness Efficient + preserves spatial structure
```

Convolution achieves dramatic parameter reduction (1000Ã— fewer!) while preserving the spatial relationships that matter for visual understanding.

## 2. Mathematical Foundations

### Understanding Convolution Step by Step

Convolution sounds complex, but it's just "sliding window multiplication and summation." Let's see exactly how it works:

```
Step 1: Position the kernel over input
Input:          Kernel:
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”     â”Œâ”€â”€â”€â”€â”€â”
â”‚ 1 2 3 4 â”‚     â”‚ 1 0 â”‚  â† Place kernel at position (0,0)
â”‚ 5 6 7 8 â”‚  Ã—  â”‚ 0 1 â”‚
â”‚ 9 0 1 2 â”‚     â””â”€â”€â”€â”€â”€â”˜
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜

Step 2: Multiply corresponding elements
Overlap:        Computation:
â”Œâ”€â”€â”€â”€â”€â”         1Ã—1 + 2Ã—0 + 5Ã—0 + 6Ã—1 = 1 + 0 + 0 + 6 = 7
â”‚ 1 2 â”‚
â”‚ 5 6 â”‚
â””â”€â”€â”€â”€â”€â”˜

Step 3: Slide kernel and repeat
Position (0,1):  Position (1,0):  Position (1,1):
â”Œâ”€â”€â”€â”€â”€â”         â”Œâ”€â”€â”€â”€â”€â”          â”Œâ”€â”€â”€â”€â”€â”
â”‚ 2 3 â”‚         â”‚ 5 6 â”‚          â”‚ 6 7 â”‚
â”‚ 6 7 â”‚         â”‚ 9 0 â”‚          â”‚ 0 1 â”‚
â””â”€â”€â”€â”€â”€â”˜         â””â”€â”€â”€â”€â”€â”˜          â””â”€â”€â”€â”€â”€â”˜
Result: 9       Result: 5        Result: 8

Final Output:   â”Œâ”€â”€â”€â”€â”€â”
               â”‚ 7 9 â”‚
               â”‚ 5 8 â”‚
               â””â”€â”€â”€â”€â”€â”˜
```

### The Mathematical Formula

For 2D convolution, we slide kernel K across input I:
```
O[i,j] = Î£ Î£ I[i+m, j+n] Ã— K[m,n]
         m n
```

This formula captures the "multiply and sum" operation for each kernel position.

### Pooling: Spatial Summarization

```
Max Pooling Example (2Ã—2 window):
Input:           Output:
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”    â”Œâ”€â”€â”€â”€â”€â”
â”‚ 1 3 2 4 â”‚    â”‚ 6 8 â”‚  â† max([1,3,5,6])=6, max([2,4,7,8])=8
â”‚ 5 6 7 8 â”‚ â†’  â”‚ 9 9 â”‚  â† max([5,2,9,1])=9, max([7,4,9,3])=9
â”‚ 2 9 1 3 â”‚    â””â”€â”€â”€â”€â”€â”˜
â”‚ 0 1 9 3 â”‚
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜

Average Pooling (same window):
â”Œâ”€â”€â”€â”€â”€â”  â† avg([1,3,5,6])=3.75, avg([2,4,7,8])=5.25
â”‚3.75 5.25â”‚
â”‚2.75 5.75â”‚  â† avg([5,2,9,1])=4.25, avg([7,4,9,3])=5.75
â””â”€â”€â”€â”€â”€â”˜
```

### Why This Complexity Matters

For convolution with input (1, 3, 224, 224) and kernel (64, 3, 3, 3):
- **Operations**: 1 Ã— 64 Ã— 3 Ã— 3 Ã— 3 Ã— 224 Ã— 224 = 86.7 million multiply-adds
- **Memory**: Input (600KB) + Weights (6.9KB) + Output (12.8MB) = ~13.4MB

This is why kernel size matters enormously - a 7Ã—7 kernel would require 5.4Ã— more computation!

### Key Properties That Enable Deep Learning

**Translation Equivariance**: Move the cat â†’ detection moves the same way
**Parameter Sharing**: Same edge detector works everywhere in the image
**Local Connectivity**: Each output only looks at nearby inputs (like human vision)
**Hierarchical Features**: Early layers detect edges â†’ later layers detect objects

## 3. Implementation - Building Spatial Operations

Now we'll implement convolution step by step, using explicit loops so you can see and feel the computational complexity. This helps you understand why modern optimizations matter!

### Conv2d: Detecting Patterns with Sliding Windows

Convolution slides a small filter (kernel) across the entire input, computing weighted sums at each position. Think of it like using a template to find matching patterns everywhere in an image.

```
Convolution Visualization:
Input (4Ã—4):              Kernel (3Ã—3):           Output (2Ã—2):
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”          â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”             â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ a b c d â”‚            â”‚ k1 k2 k3â”‚             â”‚ o1  o2 â”‚
â”‚ e f g h â”‚     Ã—      â”‚ k4 k5 k6â”‚      =      â”‚ o3  o4 â”‚
â”‚ i j k l â”‚            â”‚ k7 k8 k9â”‚             â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
â”‚ m n o p â”‚            â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜

Computation Details:
o1 = aÃ—k1 + bÃ—k2 + cÃ—k3 + eÃ—k4 + fÃ—k5 + gÃ—k6 + iÃ—k7 + jÃ—k8 + kÃ—k9
o2 = bÃ—k1 + cÃ—k2 + dÃ—k3 + fÃ—k4 + gÃ—k5 + hÃ—k6 + jÃ—k7 + kÃ—k8 + lÃ—k9
o3 = eÃ—k1 + fÃ—k2 + gÃ—k3 + iÃ—k4 + jÃ—k5 + kÃ—k6 + mÃ—k7 + nÃ—k8 + oÃ—k9
o4 = fÃ—k1 + gÃ—k2 + hÃ—k3 + jÃ—k4 + kÃ—k5 + lÃ—k6 + nÃ—k7 + oÃ—k8 + pÃ—k9
```

### The Six Nested Loops of Convolution

Our implementation will use explicit loops to show exactly where the computational cost comes from:

```
for batch in range(B):          # Loop 1: Process each sample
    for out_ch in range(C_out):     # Loop 2: Generate each output channel
        for out_h in range(H_out):      # Loop 3: Each output row
            for out_w in range(W_out):      # Loop 4: Each output column
                for k_h in range(K_h):          # Loop 5: Each kernel row
                    for k_w in range(K_w):          # Loop 6: Each kernel column
                        for in_ch in range(C_in):       # Loop 7: Each input channel
                            # The actual multiply-accumulate operation
                            result += input[...] * kernel[...]
```

Total operations: B Ã— C_out Ã— H_out Ã— W_out Ã— K_h Ã— K_w Ã— C_in

For typical values (B=32, C_out=64, H_out=224, W_out=224, K_h=3, K_w=3, C_in=3):
That's 32 Ã— 64 Ã— 224 Ã— 224 Ã— 3 Ã— 3 Ã— 3 = **2.8 billion operations** per forward pass!

### Conv2d Implementation - Building the Core of Computer Vision

Conv2d is the workhorse of computer vision. It slides learned filters across images to detect patterns like edges, textures, and eventually complex objects.

#### How Conv2d Transforms Machine Learning

```
Before Conv2d (Dense Only):         After Conv2d (Spatial Aware):
Input: 32Ã—32Ã—3 = 3,072 values      Input: 32Ã—32Ã—3 structured as image
         â†“                                   â†“
Dense(3072â†’1000) = 3M params       Conv2d(3â†’16, 3Ã—3) = 448 params
         â†“                                   â†“
No spatial awareness               Preserves spatial relationships
Massive parameter count            Parameter sharing across space
```

#### Weight Initialization: He Initialization for ReLU Networks

Our Conv2d uses He initialization, specifically designed for ReLU activations:
- **Problem**: Wrong initialization â†’ vanishing/exploding gradients
- **Solution**: std = sqrt(2 / fan_in) where fan_in = channels Ã— kernel_height Ã— kernel_width
- **Why it works**: Maintains variance through ReLU nonlinearity

#### The 6-Loop Implementation Strategy

We'll implement convolution with explicit loops to show the true computational cost:

```
Nested Loop Structure:
for batch:           â† Process each sample in parallel (in practice)
  for out_channel:   â† Generate each output feature map
    for out_h:       â† Each row of output
      for out_w:     â† Each column of output
        for k_h:     â† Each row of kernel
          for k_w:   â† Each column of kernel
            for in_ch: â† Accumulate across input channels
              result += input[...] * weight[...]
```

This reveals why convolution is expensive: O(BÃ—C_outÃ—HÃ—WÃ—K_hÃ—K_wÃ—C_in) operations!

<!-- grade_id: conv2d-class, points: 0, grade: false, solution: true -->
```python

#| export

class Conv2d:
    """
    2D Convolution layer for spatial feature extraction.

    Implements convolution with explicit loops to demonstrate
    computational complexity and memory access patterns.

    Args:
        in_channels: Number of input channels
        out_channels: Number of output feature maps
        kernel_size: Size of convolution kernel (int or tuple)
        stride: Stride of convolution (default: 1)
        padding: Zero-padding added to input (default: 0)
        bias: Whether to add learnable bias (default: True)
    """

    def __init__(self, in_channels, out_channels, kernel_size, stride=1, padding=0, bias=True):
        """
        Initialize Conv2d layer with proper weight initialization.

        TODO: Complete Conv2d initialization

        APPROACH:
        1. Store hyperparameters (channels, kernel_size, stride, padding)
        2. Initialize weights using He initialization for ReLU compatibility
        3. Initialize bias (if enabled) to zeros
        4. Use proper shapes: weight (out_channels, in_channels, kernel_h, kernel_w)

        WEIGHT INITIALIZATION:
        - He init: std = sqrt(2 / (in_channels * kernel_h * kernel_w))
        - This prevents vanishing/exploding gradients with ReLU

        HINT: Convert kernel_size to tuple if it's an integer
        """
        super().__init__()

        ### BEGIN SOLUTION
        self.in_channels = in_channels
        self.out_channels = out_channels

        # Handle kernel_size as int or tuple
        if isinstance(kernel_size, int):
            self.kernel_size = (kernel_size, kernel_size)
        else:
            self.kernel_size = kernel_size

        self.stride = stride
        self.padding = padding

        # He initialization for ReLU networks
        kernel_h, kernel_w = self.kernel_size
        fan_in = in_channels * kernel_h * kernel_w
        std = np.sqrt(2.0 / fan_in)

        # Weight shape: (out_channels, in_channels, kernel_h, kernel_w)
        self.weight = Tensor(np.random.normal(0, std,
                           (out_channels, in_channels, kernel_h, kernel_w)))

        # Bias initialization
        if bias:
            self.bias = Tensor(np.zeros(out_channels))
        else:
            self.bias = None
        ### END SOLUTION

    def forward(self, x):
        """
        Forward pass through Conv2d layer.

        TODO: Implement convolution with explicit loops

        APPROACH:
        1. Extract input dimensions and validate
        2. Calculate output dimensions
        3. Apply padding if needed
        4. Implement 6 nested loops for full convolution
        5. Add bias if present

        LOOP STRUCTURE:
        for batch in range(batch_size):
            for out_ch in range(out_channels):
                for out_h in range(out_height):
                    for out_w in range(out_width):
                        for k_h in range(kernel_height):
                            for k_w in range(kernel_width):
                                for in_ch in range(in_channels):
                                    # Accumulate: out += input * weight

        EXAMPLE:
        >>> conv = Conv2d(3, 16, kernel_size=3, padding=1)
        >>> x = Tensor(np.random.randn(2, 3, 32, 32))  # batch=2, RGB, 32x32
        >>> out = conv(x)
        >>> print(out.shape)  # Should be (2, 16, 32, 32)

        HINTS:
        - Handle padding by creating padded input array
        - Watch array bounds in inner loops
        - Accumulate products for each output position
        """
        ### BEGIN SOLUTION
        # Input validation and shape extraction
        if len(x.shape) != 4:
            raise ValueError(f"Expected 4D input (batch, channels, height, width), got {x.shape}")

        batch_size, in_channels, in_height, in_width = x.shape
        out_channels = self.out_channels
        kernel_h, kernel_w = self.kernel_size

        # Calculate output dimensions
        out_height = (in_height + 2 * self.padding - kernel_h) // self.stride + 1
        out_width = (in_width + 2 * self.padding - kernel_w) // self.stride + 1

        # Apply padding if needed
        if self.padding > 0:
            padded_input = np.pad(x.data,
                                ((0, 0), (0, 0), (self.padding, self.padding), (self.padding, self.padding)),
                                mode='constant', constant_values=0)
        else:
            padded_input = x.data

        # Initialize output
        output = np.zeros((batch_size, out_channels, out_height, out_width))

        # Explicit 6-nested loop convolution to show complexity
        for b in range(batch_size):
            for out_ch in range(out_channels):
                for out_h in range(out_height):
                    for out_w in range(out_width):
                        # Calculate input region for this output position
                        in_h_start = out_h * self.stride
                        in_w_start = out_w * self.stride

                        # Accumulate convolution result
                        conv_sum = 0.0
                        for k_h in range(kernel_h):
                            for k_w in range(kernel_w):
                                for in_ch in range(in_channels):
                                    # Get input and weight values
                                    input_val = padded_input[b, in_ch,
                                                           in_h_start + k_h,
                                                           in_w_start + k_w]
                                    weight_val = self.weight.data[out_ch, in_ch, k_h, k_w]

                                    # Accumulate
                                    conv_sum += input_val * weight_val

                        # Store result
                        output[b, out_ch, out_h, out_w] = conv_sum

        # Add bias if present
        if self.bias is not None:
            # Broadcast bias across spatial dimensions
            for out_ch in range(out_channels):
                output[:, out_ch, :, :] += self.bias.data[out_ch]

        return Tensor(output)
        ### END SOLUTION

    def parameters(self):
        """Return trainable parameters."""
        params = [self.weight]
        if self.bias is not None:
            params.append(self.bias)
        return params

    def __call__(self, x):
        """Enable model(x) syntax."""
        return self.forward(x)
```

### ðŸ§ª Unit Test: Conv2d Implementation
This test validates our convolution implementation with different configurations.
**What we're testing**: Shape preservation, padding, stride effects
**Why it matters**: Convolution is the foundation of computer vision
**Expected**: Correct output shapes and reasonable value ranges

<!-- grade_id: test-conv2d, points: 15, grade: true, solution: false -->
```python


def test_unit_conv2d():
    """ðŸ”¬ Test Conv2d implementation with multiple configurations."""
    print("ðŸ”¬ Unit Test: Conv2d...")

    # Test 1: Basic convolution without padding
    print("  Testing basic convolution...")
    conv1 = Conv2d(in_channels=3, out_channels=16, kernel_size=3)
    x1 = Tensor(np.random.randn(2, 3, 32, 32))
    out1 = conv1(x1)

    expected_h = (32 - 3) + 1  # 30
    expected_w = (32 - 3) + 1  # 30
    assert out1.shape == (2, 16, expected_h, expected_w), f"Expected (2, 16, 30, 30), got {out1.shape}"

    # Test 2: Convolution with padding (same size)
    print("  Testing convolution with padding...")
    conv2 = Conv2d(in_channels=3, out_channels=8, kernel_size=3, padding=1)
    x2 = Tensor(np.random.randn(1, 3, 28, 28))
    out2 = conv2(x2)

    # With padding=1, output should be same size as input
    assert out2.shape == (1, 8, 28, 28), f"Expected (1, 8, 28, 28), got {out2.shape}"

    # Test 3: Convolution with stride
    print("  Testing convolution with stride...")
    conv3 = Conv2d(in_channels=1, out_channels=4, kernel_size=3, stride=2)
    x3 = Tensor(np.random.randn(1, 1, 16, 16))
    out3 = conv3(x3)

    expected_h = (16 - 3) // 2 + 1  # 7
    expected_w = (16 - 3) // 2 + 1  # 7
    assert out3.shape == (1, 4, expected_h, expected_w), f"Expected (1, 4, 7, 7), got {out3.shape}"

    # Test 4: Parameter counting
    print("  Testing parameter counting...")
    conv4 = Conv2d(in_channels=64, out_channels=128, kernel_size=3, bias=True)
    params = conv4.parameters()

    # Weight: (128, 64, 3, 3) = 73,728 parameters
    # Bias: (128,) = 128 parameters
    # Total: 73,856 parameters
    weight_params = 128 * 64 * 3 * 3
    bias_params = 128
    total_params = weight_params + bias_params

    actual_weight_params = np.prod(conv4.weight.shape)
    actual_bias_params = np.prod(conv4.bias.shape) if conv4.bias is not None else 0
    actual_total = actual_weight_params + actual_bias_params

    assert actual_total == total_params, f"Expected {total_params} parameters, got {actual_total}"
    assert len(params) == 2, f"Expected 2 parameter tensors, got {len(params)}"

    # Test 5: No bias configuration
    print("  Testing no bias configuration...")
    conv5 = Conv2d(in_channels=3, out_channels=16, kernel_size=5, bias=False)
    params5 = conv5.parameters()
    assert len(params5) == 1, f"Expected 1 parameter tensor (no bias), got {len(params5)}"
    assert conv5.bias is None, "Bias should be None when bias=False"

    print("âœ… Conv2d works correctly!")

if __name__ == "__main__":
    test_unit_conv2d()
```

## 4. Pooling Operations - Spatial Dimension Reduction

Pooling operations compress spatial information while keeping the most important features. Think of them as creating "thumbnail summaries" of local regions.

### MaxPool2d: Keeping the Strongest Signals

Max pooling finds the strongest activation in each window, preserving sharp features like edges and corners.

```
MaxPool2d Example (2Ã—2 kernel, stride=2):
Input (4Ã—4):              Windows:               Output (2Ã—2):
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”          â”Œâ”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”          â”Œâ”€â”€â”€â”€â”€â”
â”‚ 1  3 â”‚ 2  8 â”‚          â”‚ 1 3 â”‚ 2 8 â”‚          â”‚ 6 8 â”‚
â”‚ 5  6 â”‚ 7  4 â”‚     â†’   â”‚ 5 6 â”‚ 7 4 â”‚    â†’    â”‚ 9 7 â”‚
â”œâ”€â”€â”€â”€â”€â”¼â”€â”€â”€â”€â”€â”¤          â”œâ”€â”€â”€â”€â”€â”¼â”€â”€â”€â”€â”€â”¤          â””â”€â”€â”€â”€â”€â”˜
â”‚ 2  9 â”‚ 1  7 â”‚          â”‚ 2 9 â”‚ 1 7 â”‚
â”‚ 0  1 â”‚ 3  6 â”‚          â”‚ 0 1 â”‚ 3 6 â”‚
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜          â””â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”˜

Window Computations:
Top-left: max(1,3,5,6) = 6     Top-right: max(2,8,7,4) = 8
Bottom-left: max(2,9,0,1) = 9  Bottom-right: max(1,7,3,6) = 7
```

### AvgPool2d: Smoothing Local Features

Average pooling computes the mean of each window, creating smoother, more general features.

```
AvgPool2d Example (same 2Ã—2 kernel, stride=2):
Input (4Ã—4):              Output (2Ã—2):
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”          â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ 1  3 â”‚ 2  8 â”‚          â”‚ 3.75  5.25â”‚
â”‚ 5  6 â”‚ 7  4 â”‚     â†’   â”‚ 3.0   4.25â”‚
â”œâ”€â”€â”€â”€â”€â”¼â”€â”€â”€â”€â”€â”¤          â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
â”‚ 2  9 â”‚ 1  7 â”‚
â”‚ 0  1 â”‚ 3  6 â”‚
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜

Window Computations:
Top-left: (1+3+5+6)/4 = 3.75    Top-right: (2+8+7+4)/4 = 5.25
Bottom-left: (2+9+0+1)/4 = 3.0  Bottom-right: (1+7+3+6)/4 = 4.25
```

### Why Pooling Matters for Computer Vision

```
Memory Impact:
Input: 224Ã—224Ã—64 = 3.2M values    After 2Ã—2 pooling: 112Ã—112Ã—64 = 0.8M values
Memory reduction: 4Ã— less!         Computation reduction: 4Ã— less!

Information Trade-off:
âœ… Preserves important features     âš ï¸ Loses fine spatial detail
âœ… Provides translation invariance  âš ï¸ Reduces localization precision
âœ… Reduces overfitting             âš ï¸ May lose small objects
```

### Sliding Window Pattern

Both pooling operations follow the same sliding window pattern:

```
Sliding 2Ã—2 window with stride=2:
Step 1:     Step 2:     Step 3:     Step 4:
â”Œâ”€â”€â”        â”Œâ”€â”€â”
â”‚â–“â–“â”‚        â”‚â–“â–“â”‚
â””â”€â”€â”˜        â””â”€â”€â”˜                   â”Œâ”€â”€â”        â”Œâ”€â”€â”
                                    â”‚â–“â–“â”‚        â”‚â–“â–“â”‚
                                    â””â”€â”€â”˜        â””â”€â”€â”˜

Non-overlapping windows â†’ Each input pixel used exactly once
Stride=2 â†’ Output dimensions halved in each direction
```

The key difference: MaxPool takes max(window), AvgPool takes mean(window).

### MaxPool2d Implementation - Preserving Strong Features

MaxPool2d finds the strongest activation in each spatial window, creating a compressed representation that keeps the most important information.

#### Why Max Pooling Works for Computer Vision

```
Edge Detection Example:
Input Window (2Ã—2):         Max Pooling Result:
â”Œâ”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”
â”‚ 0.1 â”‚ 0.8 â”‚ â†  Strong edge signal
â”œâ”€â”€â”€â”€â”€â”¼â”€â”€â”€â”€â”€â”¤
â”‚ 0.2 â”‚ 0.1 â”‚              Output: 0.8 (preserves edge)
â””â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”˜

Noise Reduction Example:
Input Window (2Ã—2):
â”Œâ”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”
â”‚ 0.9 â”‚ 0.1 â”‚ â†  Feature + noise
â”œâ”€â”€â”€â”€â”€â”¼â”€â”€â”€â”€â”€â”¤
â”‚ 0.2 â”‚ 0.1 â”‚              Output: 0.9 (removes noise)
â””â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”˜
```

#### The Sliding Window Pattern

```
MaxPool with 2Ã—2 kernel, stride=2:

Input (4Ã—4):                Output (2Ã—2):
â”Œâ”€â”€â”€â”¬â”€â”€â”€â”¬â”€â”€â”€â”¬â”€â”€â”€â”          â”Œâ”€â”€â”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”€â”€â”
â”‚ a â”‚ b â”‚ c â”‚ d â”‚          â”‚max(a,bâ”‚max(c,dâ”‚
â”œâ”€â”€â”€â”¼â”€â”€â”€â”¼â”€â”€â”€â”¼â”€â”€â”€â”¤     â†’    â”‚   e,f)â”‚   g,h)â”‚
â”‚ e â”‚ f â”‚ g â”‚ h â”‚          â”œâ”€â”€â”€â”€â”€â”€â”€â”¼â”€â”€â”€â”€â”€â”€â”€â”¤
â”œâ”€â”€â”€â”¼â”€â”€â”€â”¼â”€â”€â”€â”¼â”€â”€â”€â”¤          â”‚max(i,jâ”‚max(k,lâ”‚
â”‚ i â”‚ j â”‚ k â”‚ l â”‚          â”‚   m,n)â”‚   o,p)â”‚
â”œâ”€â”€â”€â”¼â”€â”€â”€â”¼â”€â”€â”€â”¼â”€â”€â”€â”¤          â””â”€â”€â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”€â”€â”˜
â”‚ m â”‚ n â”‚ o â”‚ p â”‚
â””â”€â”€â”€â”´â”€â”€â”€â”´â”€â”€â”€â”´â”€â”€â”€â”˜

Benefits:
âœ“ Translation invariance (cat moved 1 pixel still detected)
âœ“ Computational efficiency (4Ã— fewer values to process)
âœ“ Hierarchical feature building (next layer sees larger receptive field)
```

#### Memory and Computation Impact

For input (1, 64, 224, 224) with 2Ã—2 pooling:
- **Input memory**: 64 Ã— 224 Ã— 224 Ã— 4 bytes = 12.8 MB
- **Output memory**: 64 Ã— 112 Ã— 112 Ã— 4 bytes = 3.2 MB
- **Memory reduction**: 4Ã— less memory needed
- **Computation**: No parameters, minimal compute cost

<!-- grade_id: maxpool2d-class, points: 0, grade: false, solution: true -->
```python

#| export

class MaxPool2d:
    """
    2D Max Pooling layer for spatial dimension reduction.

    Applies maximum operation over spatial windows, preserving
    the strongest activations while reducing computational load.

    Args:
        kernel_size: Size of pooling window (int or tuple)
        stride: Stride of pooling operation (default: same as kernel_size)
        padding: Zero-padding added to input (default: 0)
    """

    def __init__(self, kernel_size, stride=None, padding=0):
        """
        Initialize MaxPool2d layer.

        TODO: Store pooling parameters

        APPROACH:
        1. Convert kernel_size to tuple if needed
        2. Set stride to kernel_size if not provided (non-overlapping)
        3. Store padding parameter

        HINT: Default stride equals kernel_size for non-overlapping windows
        """
        super().__init__()

        ### BEGIN SOLUTION
        # Handle kernel_size as int or tuple
        if isinstance(kernel_size, int):
            self.kernel_size = (kernel_size, kernel_size)
        else:
            self.kernel_size = kernel_size

        # Default stride equals kernel_size (non-overlapping)
        if stride is None:
            self.stride = self.kernel_size[0]
        else:
            self.stride = stride

        self.padding = padding
        ### END SOLUTION

    def forward(self, x):
        """
        Forward pass through MaxPool2d layer.

        TODO: Implement max pooling with explicit loops

        APPROACH:
        1. Extract input dimensions
        2. Calculate output dimensions
        3. Apply padding if needed
        4. Implement nested loops for pooling windows
        5. Find maximum value in each window

        LOOP STRUCTURE:
        for batch in range(batch_size):
            for channel in range(channels):
                for out_h in range(out_height):
                    for out_w in range(out_width):
                        # Find max in window [in_h:in_h+k_h, in_w:in_w+k_w]
                        max_val = -infinity
                        for k_h in range(kernel_height):
                            for k_w in range(kernel_width):
                                max_val = max(max_val, input[...])

        EXAMPLE:
        >>> pool = MaxPool2d(kernel_size=2, stride=2)
        >>> x = Tensor(np.random.randn(1, 3, 8, 8))
        >>> out = pool(x)
        >>> print(out.shape)  # Should be (1, 3, 4, 4)

        HINTS:
        - Initialize max_val to negative infinity
        - Handle stride correctly when accessing input
        - No parameters to update (pooling has no weights)
        """
        ### BEGIN SOLUTION
        # Input validation and shape extraction
        if len(x.shape) != 4:
            raise ValueError(f"Expected 4D input (batch, channels, height, width), got {x.shape}")

        batch_size, channels, in_height, in_width = x.shape
        kernel_h, kernel_w = self.kernel_size

        # Calculate output dimensions
        out_height = (in_height + 2 * self.padding - kernel_h) // self.stride + 1
        out_width = (in_width + 2 * self.padding - kernel_w) // self.stride + 1

        # Apply padding if needed
        if self.padding > 0:
            padded_input = np.pad(x.data,
                                ((0, 0), (0, 0), (self.padding, self.padding), (self.padding, self.padding)),
                                mode='constant', constant_values=-np.inf)
        else:
            padded_input = x.data

        # Initialize output
        output = np.zeros((batch_size, channels, out_height, out_width))

        # Explicit nested loop max pooling
        for b in range(batch_size):
            for c in range(channels):
                for out_h in range(out_height):
                    for out_w in range(out_width):
                        # Calculate input region for this output position
                        in_h_start = out_h * self.stride
                        in_w_start = out_w * self.stride

                        # Find maximum in window
                        max_val = -np.inf
                        for k_h in range(kernel_h):
                            for k_w in range(kernel_w):
                                input_val = padded_input[b, c,
                                                       in_h_start + k_h,
                                                       in_w_start + k_w]
                                max_val = max(max_val, input_val)

                        # Store result
                        output[b, c, out_h, out_w] = max_val

        return Tensor(output)
        ### END SOLUTION

    def parameters(self):
        """Return empty list (pooling has no parameters)."""
        return []

    def __call__(self, x):
        """Enable model(x) syntax."""
        return self.forward(x)
```

### AvgPool2d Implementation - Smoothing and Generalizing Features

AvgPool2d computes the average of each spatial window, creating smoother features that are less sensitive to noise and exact pixel positions.

#### MaxPool vs AvgPool: Different Philosophies

```
Same Input Window (2Ã—2):    MaxPool Output:    AvgPool Output:
â”Œâ”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”
â”‚ 0.1 â”‚ 0.9 â”‚               0.9              0.425
â”œâ”€â”€â”€â”€â”€â”¼â”€â”€â”€â”€â”€â”¤              (max)             (mean)
â”‚ 0.3 â”‚ 0.3 â”‚
â””â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”˜

Interpretation:
MaxPool: "What's the strongest feature here?"
AvgPool: "What's the general feature level here?"
```

#### When to Use Average Pooling

```
Use Cases:
âœ“ Global Average Pooling (GAP) for classification
âœ“ When you want smoother, less noisy features
âœ“ When exact feature location doesn't matter
âœ“ In shallower networks where sharp features aren't critical

Typical Pattern:
Feature Maps â†’ Global Average Pool â†’ Dense â†’ Classification
(256Ã—7Ã—7)   â†’        (256Ã—1Ã—1)      â†’ FC   â†’    (10)
              Replaces flatten+dense with parameter reduction
```

#### Mathematical Implementation

```
Average Pooling Computation:
Window: [a, b]    Result = (a + b + c + d) / 4
        [c, d]

For efficiency, we:
1. Sum all values in window: window_sum = a + b + c + d
2. Divide by window area: result = window_sum / (kernel_h Ã— kernel_w)
3. Store result at output position

Memory access pattern identical to MaxPool, just different aggregation!
```

#### Practical Considerations

- **Memory**: Same 4Ã— reduction as MaxPool
- **Computation**: Slightly more expensive (sum + divide vs max)
- **Features**: Smoother, more generalized than MaxPool
- **Use**: Often in final layers (Global Average Pooling) to reduce parameters

<!-- grade_id: avgpool2d-class, points: 0, grade: false, solution: true -->
```python

#| export

class AvgPool2d:
    """
    2D Average Pooling layer for spatial dimension reduction.

    Applies average operation over spatial windows, smoothing
    features while reducing computational load.

    Args:
        kernel_size: Size of pooling window (int or tuple)
        stride: Stride of pooling operation (default: same as kernel_size)
        padding: Zero-padding added to input (default: 0)
    """

    def __init__(self, kernel_size, stride=None, padding=0):
        """
        Initialize AvgPool2d layer.

        TODO: Store pooling parameters (same as MaxPool2d)

        APPROACH:
        1. Convert kernel_size to tuple if needed
        2. Set stride to kernel_size if not provided
        3. Store padding parameter
        """
        super().__init__()

        ### BEGIN SOLUTION
        # Handle kernel_size as int or tuple
        if isinstance(kernel_size, int):
            self.kernel_size = (kernel_size, kernel_size)
        else:
            self.kernel_size = kernel_size

        # Default stride equals kernel_size (non-overlapping)
        if stride is None:
            self.stride = self.kernel_size[0]
        else:
            self.stride = stride

        self.padding = padding
        ### END SOLUTION

    def forward(self, x):
        """
        Forward pass through AvgPool2d layer.

        TODO: Implement average pooling with explicit loops

        APPROACH:
        1. Similar structure to MaxPool2d
        2. Instead of max, compute average of window
        3. Divide sum by window area for true average

        LOOP STRUCTURE:
        for batch in range(batch_size):
            for channel in range(channels):
                for out_h in range(out_height):
                    for out_w in range(out_width):
                        # Compute average in window
                        window_sum = 0
                        for k_h in range(kernel_height):
                            for k_w in range(kernel_width):
                                window_sum += input[...]
                        avg_val = window_sum / (kernel_height * kernel_width)

        HINT: Remember to divide by window area to get true average
        """
        ### BEGIN SOLUTION
        # Input validation and shape extraction
        if len(x.shape) != 4:
            raise ValueError(f"Expected 4D input (batch, channels, height, width), got {x.shape}")

        batch_size, channels, in_height, in_width = x.shape
        kernel_h, kernel_w = self.kernel_size

        # Calculate output dimensions
        out_height = (in_height + 2 * self.padding - kernel_h) // self.stride + 1
        out_width = (in_width + 2 * self.padding - kernel_w) // self.stride + 1

        # Apply padding if needed
        if self.padding > 0:
            padded_input = np.pad(x.data,
                                ((0, 0), (0, 0), (self.padding, self.padding), (self.padding, self.padding)),
                                mode='constant', constant_values=0)
        else:
            padded_input = x.data

        # Initialize output
        output = np.zeros((batch_size, channels, out_height, out_width))

        # Explicit nested loop average pooling
        for b in range(batch_size):
            for c in range(channels):
                for out_h in range(out_height):
                    for out_w in range(out_width):
                        # Calculate input region for this output position
                        in_h_start = out_h * self.stride
                        in_w_start = out_w * self.stride

                        # Compute sum in window
                        window_sum = 0.0
                        for k_h in range(kernel_h):
                            for k_w in range(kernel_w):
                                input_val = padded_input[b, c,
                                                       in_h_start + k_h,
                                                       in_w_start + k_w]
                                window_sum += input_val

                        # Compute average
                        avg_val = window_sum / (kernel_h * kernel_w)

                        # Store result
                        output[b, c, out_h, out_w] = avg_val

        return Tensor(output)
        ### END SOLUTION

    def parameters(self):
        """Return empty list (pooling has no parameters)."""
        return []

    def __call__(self, x):
        """Enable model(x) syntax."""
        return self.forward(x)
```

### ðŸ§ª Unit Test: Pooling Operations
This test validates both max and average pooling implementations.
**What we're testing**: Dimension reduction, aggregation correctness
**Why it matters**: Pooling is essential for computational efficiency in CNNs
**Expected**: Correct output shapes and proper value aggregation

<!-- grade_id: test-pooling, points: 10, grade: true, solution: false -->
```python


def test_unit_pooling():
    """ðŸ”¬ Test MaxPool2d and AvgPool2d implementations."""
    print("ðŸ”¬ Unit Test: Pooling Operations...")

    # Test 1: MaxPool2d basic functionality
    print("  Testing MaxPool2d...")
    maxpool = MaxPool2d(kernel_size=2, stride=2)
    x1 = Tensor(np.random.randn(1, 3, 8, 8))
    out1 = maxpool(x1)

    expected_shape = (1, 3, 4, 4)  # 8/2 = 4
    assert out1.shape == expected_shape, f"MaxPool expected {expected_shape}, got {out1.shape}"

    # Test 2: AvgPool2d basic functionality
    print("  Testing AvgPool2d...")
    avgpool = AvgPool2d(kernel_size=2, stride=2)
    x2 = Tensor(np.random.randn(2, 16, 16, 16))
    out2 = avgpool(x2)

    expected_shape = (2, 16, 8, 8)  # 16/2 = 8
    assert out2.shape == expected_shape, f"AvgPool expected {expected_shape}, got {out2.shape}"

    # Test 3: MaxPool vs AvgPool on known data
    print("  Testing max vs avg behavior...")
    # Create simple test case with known values
    test_data = np.array([[[[1, 2, 3, 4],
                           [5, 6, 7, 8],
                           [9, 10, 11, 12],
                           [13, 14, 15, 16]]]], dtype=np.float32)
    x3 = Tensor(test_data)

    maxpool_test = MaxPool2d(kernel_size=2, stride=2)
    avgpool_test = AvgPool2d(kernel_size=2, stride=2)

    max_out = maxpool_test(x3)
    avg_out = avgpool_test(x3)

    # For 2x2 windows:
    # Top-left: max([1,2,5,6]) = 6, avg = 3.5
    # Top-right: max([3,4,7,8]) = 8, avg = 5.5
    # Bottom-left: max([9,10,13,14]) = 14, avg = 11.5
    # Bottom-right: max([11,12,15,16]) = 16, avg = 13.5

    expected_max = np.array([[[[6, 8], [14, 16]]]])
    expected_avg = np.array([[[[3.5, 5.5], [11.5, 13.5]]]])

    assert np.allclose(max_out.data, expected_max), f"MaxPool values incorrect: {max_out.data} vs {expected_max}"
    assert np.allclose(avg_out.data, expected_avg), f"AvgPool values incorrect: {avg_out.data} vs {expected_avg}"

    # Test 4: Overlapping pooling (stride < kernel_size)
    print("  Testing overlapping pooling...")
    overlap_pool = MaxPool2d(kernel_size=3, stride=1)
    x4 = Tensor(np.random.randn(1, 1, 5, 5))
    out4 = overlap_pool(x4)

    # Output: (5-3)/1 + 1 = 3
    expected_shape = (1, 1, 3, 3)
    assert out4.shape == expected_shape, f"Overlapping pool expected {expected_shape}, got {out4.shape}"

    # Test 5: No parameters in pooling layers
    print("  Testing parameter counts...")
    assert len(maxpool.parameters()) == 0, "MaxPool should have no parameters"
    assert len(avgpool.parameters()) == 0, "AvgPool should have no parameters"

    print("âœ… Pooling operations work correctly!")

if __name__ == "__main__":
    test_unit_pooling()
```

## 5. Systems Analysis - Understanding Spatial Operation Performance

Now let's analyze the computational complexity and memory trade-offs of spatial operations. This analysis reveals why certain design choices matter for real-world performance.

### Key Questions We'll Answer:
1. How does convolution complexity scale with input size and kernel size?
2. What's the memory vs computation trade-off in different approaches?
3. How do modern optimizations (like im2col) change the performance characteristics?

<!-- grade_id: spatial-analysis, points: 0, grade: false, solution: true -->
```python


def analyze_convolution_complexity():
    """ðŸ“Š Analyze convolution computational complexity across different configurations."""
    print("ðŸ“Š Analyzing Convolution Complexity...")

    # Test configurations optimized for educational demonstration (smaller sizes)
    configs = [
        {"input": (1, 3, 16, 16), "conv": (8, 3, 3), "name": "Small (16Ã—16)"},
        {"input": (1, 3, 24, 24), "conv": (12, 3, 3), "name": "Medium (24Ã—24)"},
        {"input": (1, 3, 32, 32), "conv": (16, 3, 3), "name": "Large (32Ã—32)"},
        {"input": (1, 3, 16, 16), "conv": (8, 3, 5), "name": "Large Kernel (5Ã—5)"},
    ]

    print(f"{'Configuration':<20} {'FLOPs':<15} {'Memory (MB)':<12} {'Time (ms)':<10}")
    print("-" * 70)

    for config in configs:
        # Create convolution layer
        in_ch = config["input"][1]
        out_ch, k_size = config["conv"][0], config["conv"][1]
        conv = Conv2d(in_ch, out_ch, kernel_size=k_size, padding=k_size//2)

        # Create input tensor
        x = Tensor(np.random.randn(*config["input"]))

        # Calculate theoretical FLOPs
        batch, in_channels, h, w = config["input"]
        out_channels, kernel_size = config["conv"][0], config["conv"][1]

        # Each output element requires in_channels * kernel_sizeÂ² multiply-adds
        flops_per_output = in_channels * kernel_size * kernel_size * 2  # 2 for MAC
        total_outputs = batch * out_channels * h * w  # Assuming same size with padding
        total_flops = flops_per_output * total_outputs

        # Measure memory usage
        input_memory = np.prod(config["input"]) * 4  # float32 = 4 bytes
        weight_memory = out_channels * in_channels * kernel_size * kernel_size * 4
        output_memory = batch * out_channels * h * w * 4
        total_memory = (input_memory + weight_memory + output_memory) / (1024 * 1024)  # MB

        # Measure execution time
        start_time = time.time()
        _ = conv(x)
        end_time = time.time()
        exec_time = (end_time - start_time) * 1000  # ms

        print(f"{config['name']:<20} {total_flops:<15,} {total_memory:<12.2f} {exec_time:<10.2f}")

    print("\nðŸ’¡ Key Insights:")
    print("ðŸ”¸ FLOPs scale as O(HÃ—WÃ—C_inÃ—C_outÃ—KÂ²) - quadratic in spatial and kernel size")
    print("ðŸ”¸ Memory scales linearly with spatial dimensions and channels")
    print("ðŸ”¸ Large kernels dramatically increase computational cost")
    print("ðŸš€ This motivates depthwise separable convolutions and attention mechanisms")

# Analysis will be called in main execution
```

<!-- grade_id: pooling-analysis, points: 0, grade: false, solution: true -->
```python


def analyze_pooling_effects():
    """ðŸ“Š Analyze pooling's impact on spatial dimensions and features."""
    print("\nðŸ“Š Analyzing Pooling Effects...")

    # Create sample input with spatial structure
    # Simple edge pattern that pooling should preserve differently
    pattern = np.zeros((1, 1, 8, 8))
    pattern[0, 0, :, 3:5] = 1.0  # Vertical edge
    pattern[0, 0, 3:5, :] = 1.0  # Horizontal edge
    x = Tensor(pattern)

    print("Original 8Ã—8 pattern:")
    print(x.data[0, 0])

    # Test different pooling strategies
    pools = [
        (MaxPool2d(2, stride=2), "MaxPool 2Ã—2"),
        (AvgPool2d(2, stride=2), "AvgPool 2Ã—2"),
        (MaxPool2d(4, stride=4), "MaxPool 4Ã—4"),
        (AvgPool2d(4, stride=4), "AvgPool 4Ã—4"),
    ]

    print(f"\n{'Operation':<15} {'Output Shape':<15} {'Feature Preservation'}")
    print("-" * 60)

    for pool_op, name in pools:
        result = pool_op(x)
        # Measure how much of the original pattern is preserved
        preservation = np.sum(result.data > 0.1) / np.prod(result.shape)
        print(f"{name:<15} {str(result.shape):<15} {preservation:<.2%}")

        print(f"  Output:")
        print(f"  {result.data[0, 0]}")
        print()

    print("ðŸ’¡ Key Insights:")
    print("ðŸ”¸ MaxPool preserves sharp features better (edge detection)")
    print("ðŸ”¸ AvgPool smooths features (noise reduction)")
    print("ðŸ”¸ Larger pooling windows lose more spatial detail")
    print("ðŸš€ Choice depends on task: classification vs detection vs segmentation")

# Analysis will be called in main execution
```

## 6. Integration - Building a Complete CNN

Now let's combine convolution and pooling into a complete CNN architecture. You'll see how spatial operations work together to transform raw pixels into meaningful features.

### CNN Architecture: From Pixels to Predictions

A CNN processes images through alternating convolution and pooling layers, gradually extracting higher-level features:

```
Complete CNN Pipeline:

Input Image (32Ã—32Ã—3)     Raw RGB pixels
       â†“
Conv2d(3â†’16, 3Ã—3)        Detect edges, textures
       â†“
ReLU Activation          Remove negative values
       â†“
MaxPool(2Ã—2)             Reduce to (16Ã—16Ã—16)
       â†“
Conv2d(16â†’32, 3Ã—3)       Detect shapes, patterns
       â†“
ReLU Activation          Remove negative values
       â†“
MaxPool(2Ã—2)             Reduce to (8Ã—8Ã—32)
       â†“
Flatten                  Reshape to vector (2048,)
       â†“
Linear(2048â†’10)          Final classification
       â†“
Softmax                  Probability distribution
```

### The Parameter Efficiency Story

```
CNN vs Dense Network Comparison:

CNN Approach:                     Dense Approach:
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”               â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ Conv1: 3â†’16     â”‚               â”‚ Input: 32Ã—32Ã—3  â”‚
â”‚ Params: 448     â”‚               â”‚ = 3,072 values  â”‚
â”œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”¤               â”œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”¤
â”‚ Conv2: 16â†’32    â”‚               â”‚ Hidden: 1,000   â”‚
â”‚ Params: 4,640   â”‚               â”‚ Params: 3M+     â”‚
â”œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”¤               â”œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”¤
â”‚ Linear: 2048â†’10 â”‚               â”‚ Output: 10      â”‚
â”‚ Params: 20,490  â”‚               â”‚ Params: 10K     â”‚
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜               â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
Total: ~25K params                Total: ~3M params

CNN wins with 120Ã— fewer parameters!
```

### Spatial Hierarchy: Why This Architecture Works

```
Layer-by-Layer Feature Evolution:

Layer 1 (Conv 3â†’16):              Layer 2 (Conv 16â†’32):
â”Œâ”€â”€â”€â”€â”€â” â”Œâ”€â”€â”€â”€â”€â” â”Œâ”€â”€â”€â”€â”€â”           â”Œâ”€â”€â”€â”€â”€â” â”Œâ”€â”€â”€â”€â”€â” â”Œâ”€â”€â”€â”€â”€â”
â”‚Edge â”‚ â”‚Edge â”‚ â”‚Edge â”‚           â”‚Shapeâ”‚ â”‚Cornerâ”‚ â”‚Textureâ”‚
â”‚ \\ /â”‚ â”‚  |  â”‚ â”‚ / \\â”‚           â”‚ â—‡  â”‚ â”‚  L  â”‚ â”‚ â‰ˆâ‰ˆâ‰ˆ â”‚
â””â”€â”€â”€â”€â”€â”˜ â””â”€â”€â”€â”€â”€â”˜ â””â”€â”€â”€â”€â”€â”˜           â””â”€â”€â”€â”€â”€â”˜ â””â”€â”€â”€â”€â”€â”˜ â””â”€â”€â”€â”€â”€â”˜
Simple features                   Complex combinations

Why pooling between layers:
âœ“ Reduces computation for next layer
âœ“ Increases receptive field (each conv sees larger input area)
âœ“ Provides translation invariance (cat moved 1 pixel still detected)
```

This hierarchical approach mirrors human vision: we first detect edges, then shapes, then objects!

### SimpleCNN Implementation - Putting It All Together

Now we'll build a complete CNN that demonstrates how convolution and pooling work together. This is your first step from processing individual tensors to understanding complete images!

#### The CNN Architecture Pattern

```
SimpleCNN Architecture Visualization:

Input: (batch, 3, 32, 32)     â† RGB images (CIFAR-10 size)
         â†“
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ Conv2d(3â†’16, 3Ã—3, p=1) â”‚    â† Detect edges, textures
â”‚ ReLU()                  â”‚    â† Remove negative values
â”‚ MaxPool(2Ã—2)            â”‚    â† Reduce to (batch, 16, 16, 16)
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
         â†“
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ Conv2d(16â†’32, 3Ã—3, p=1) â”‚   â† Detect shapes, patterns
â”‚ ReLU()                  â”‚   â† Remove negative values
â”‚ MaxPool(2Ã—2)            â”‚   â† Reduce to (batch, 32, 8, 8)
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
         â†“
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ Flatten()               â”‚   â† Reshape to (batch, 2048)
â”‚ Linear(2048â†’10)         â”‚   â† Final classification
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
         â†“
Output: (batch, 10)           â† Class probabilities
```

#### Why This Architecture Works

```
Feature Hierarchy Development:

Layer 1 Features (3â†’16):     Layer 2 Features (16â†’32):
â”Œâ”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”   â”Œâ”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”¬â”€â”€â”€â”€â”€â”
â”‚Edge â”‚Edge â”‚Edge â”‚Blob â”‚   â”‚Shapeâ”‚Cornerâ”‚Tex-â”‚Pat- â”‚
â”‚ \\  â”‚  |  â”‚ /   â”‚  â—‹  â”‚   â”‚ â—‡   â”‚  L  â”‚tureâ”‚tern â”‚
â””â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”˜   â””â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”´â”€â”€â”€â”€â”€â”˜
Simple features             Complex combinations

Spatial Dimension Reduction:
32Ã—32 â†’ 16Ã—16 â†’ 8Ã—8
 1024    256     64  (per channel)

Channel Expansion:
3 â†’ 16 â†’ 32
More feature types at each level
```

#### Parameter Efficiency Demonstration

```
CNN vs Dense Comparison for 32Ã—32Ã—3 â†’ 10 classes:

CNN Approach:                    Dense Approach:
â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”          â”Œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”
â”‚ Conv1: 3â†’16, 3Ã—3   â”‚          â”‚ Input: 3072 values â”‚
â”‚ Params: 448        â”‚          â”‚        â†“          â”‚
â”œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”¤          â”‚ Dense: 3072â†’512   â”‚
â”‚ Conv2: 16â†’32, 3Ã—3  â”‚          â”‚ Params: 1.57M     â”‚
â”‚ Params: 4,640      â”‚          â”œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”¤
â”œâ”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”¤          â”‚ Dense: 512â†’10     â”‚
â”‚ Dense: 2048â†’10     â”‚          â”‚ Params: 5,120     â”‚
â”‚ Params: 20,490     â”‚          â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜
â””â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”€â”˜          Total: 1.58M params
Total: 25,578 params

CNN has 62Ã— fewer parameters while preserving spatial structure!
```

#### Receptive Field Growth

```
How each layer sees progressively larger input regions:

Layer 1 Conv (3Ã—3):           Layer 2 Conv (3Ã—3):
Each output pixel sees        Each output pixel sees
3Ã—3 = 9 input pixels         7Ã—7 = 49 input pixels
                             (due to pooling+conv)

Final Result: Layer 2 can detect complex patterns
spanning 7Ã—7 regions of original image!
```

<!-- grade_id: simple-cnn, points: 0, grade: false, solution: true -->
```python

#| export

class SimpleCNN:
    """
    Simple CNN demonstrating spatial operations integration.

    Architecture:
    - Conv2d(3â†’16, 3Ã—3) + ReLU + MaxPool(2Ã—2)
    - Conv2d(16â†’32, 3Ã—3) + ReLU + MaxPool(2Ã—2)
    - Flatten + Linear(featuresâ†’num_classes)
    """

    def __init__(self, num_classes=10):
        """
        Initialize SimpleCNN.

        TODO: Build CNN architecture with spatial and dense layers

        APPROACH:
        1. Conv layer 1: 3 â†’ 16 channels, 3Ã—3 kernel, padding=1
        2. Pool layer 1: 2Ã—2 max pooling
        3. Conv layer 2: 16 â†’ 32 channels, 3Ã—3 kernel, padding=1
        4. Pool layer 2: 2Ã—2 max pooling
        5. Calculate flattened size and add final linear layer

        HINT: For 32Ã—32 input â†’ 32â†’16â†’8â†’4 spatial reduction
        Final feature size: 32 channels Ã— 4Ã—4 = 512 features
        """
        super().__init__()

        ### BEGIN SOLUTION
        # Convolutional layers
        self.conv1 = Conv2d(in_channels=3, out_channels=16, kernel_size=3, padding=1)
        self.pool1 = MaxPool2d(kernel_size=2, stride=2)

        self.conv2 = Conv2d(in_channels=16, out_channels=32, kernel_size=3, padding=1)
        self.pool2 = MaxPool2d(kernel_size=2, stride=2)

        # Calculate flattened size
        # Input: 32Ã—32 â†’ Conv1+Pool1: 16Ã—16 â†’ Conv2+Pool2: 8Ã—8
        # Wait, let's recalculate: 32Ã—32 â†’ Pool1: 16Ã—16 â†’ Pool2: 8Ã—8
        # Final: 32 channels Ã— 8Ã—8 = 2048 features
        self.flattened_size = 32 * 8 * 8

        # Import Linear layer (we'll implement a simple version)
        # For now, we'll use a placeholder that we can replace
        # This represents the final classification layer
        self.num_classes = num_classes
        self.flattened_size = 32 * 8 * 8  # Will be used when we add Linear layer
        ### END SOLUTION

    def forward(self, x):
        """
        Forward pass through SimpleCNN.

        TODO: Implement CNN forward pass

        APPROACH:
        1. Apply conv1 â†’ ReLU â†’ pool1
        2. Apply conv2 â†’ ReLU â†’ pool2
        3. Flatten spatial dimensions
        4. Apply final linear layer (when available)

        For now, return features before final linear layer
        since we haven't imported Linear from layers module yet.
        """
        ### BEGIN SOLUTION
        # First conv block
        x = self.conv1(x)
        x = self.relu(x)  # ReLU activation
        x = self.pool1(x)

        # Second conv block
        x = self.conv2(x)
        x = self.relu(x)  # ReLU activation
        x = self.pool2(x)

        # Flatten for classification (reshape to 2D)
        batch_size = x.shape[0]
        x_flat = x.data.reshape(batch_size, -1)

        # Return flattened features
        # In a complete implementation, this would go through a Linear layer
        return Tensor(x_flat)
        ### END SOLUTION

    def relu(self, x):
        """Simple ReLU implementation for CNN."""
        return Tensor(np.maximum(0, x.data))

    def parameters(self):
        """Return all trainable parameters."""
        params = []
        params.extend(self.conv1.parameters())
        params.extend(self.conv2.parameters())
        # Linear layer parameters would be added here
        return params

    def __call__(self, x):
        """Enable model(x) syntax."""
        return self.forward(x)
```

### ðŸ§ª Unit Test: SimpleCNN Integration
This test validates that spatial operations work together in a complete CNN architecture.
**What we're testing**: End-to-end spatial processing pipeline
**Why it matters**: Spatial operations must compose correctly for real CNNs
**Expected**: Proper dimension reduction and feature extraction

<!-- grade_id: test-simple-cnn, points: 10, grade: true, solution: false -->
```python


def test_unit_simple_cnn():
    """ðŸ”¬ Test SimpleCNN integration with spatial operations."""
    print("ðŸ”¬ Unit Test: SimpleCNN Integration...")

    # Test 1: Forward pass with CIFAR-10 sized input
    print("  Testing forward pass...")
    model = SimpleCNN(num_classes=10)
    x = Tensor(np.random.randn(2, 3, 32, 32))  # Batch of 2, RGB, 32Ã—32

    features = model(x)

    # Expected: 2 samples, 32 channels Ã— 8Ã—8 spatial = 2048 features
    expected_shape = (2, 2048)
    assert features.shape == expected_shape, f"Expected {expected_shape}, got {features.shape}"

    # Test 2: Parameter counting
    print("  Testing parameter counting...")
    params = model.parameters()

    # Conv1: (16, 3, 3, 3) + bias (16,) = 432 + 16 = 448
    # Conv2: (32, 16, 3, 3) + bias (32,) = 4608 + 32 = 4640
    # Total: 448 + 4640 = 5088 parameters

    conv1_params = 16 * 3 * 3 * 3 + 16  # weights + bias
    conv2_params = 32 * 16 * 3 * 3 + 32  # weights + bias
    expected_total = conv1_params + conv2_params

    actual_total = sum(np.prod(p.shape) for p in params)
    assert actual_total == expected_total, f"Expected {expected_total} parameters, got {actual_total}"

    # Test 3: Different input sizes
    print("  Testing different input sizes...")

    # Test with different spatial dimensions
    x_small = Tensor(np.random.randn(1, 3, 16, 16))
    features_small = model(x_small)

    # 16Ã—16 â†’ 8Ã—8 â†’ 4Ã—4, so 32 Ã— 4Ã—4 = 512 features
    expected_small = (1, 512)
    assert features_small.shape == expected_small, f"Expected {expected_small}, got {features_small.shape}"

    # Test 4: Batch processing
    print("  Testing batch processing...")
    x_batch = Tensor(np.random.randn(8, 3, 32, 32))
    features_batch = model(x_batch)

    expected_batch = (8, 2048)
    assert features_batch.shape == expected_batch, f"Expected {expected_batch}, got {features_batch.shape}"

    print("âœ… SimpleCNN integration works correctly!")

if __name__ == "__main__":
    test_unit_simple_cnn()
```

## 7. Module Integration Test

Final validation that everything works together correctly.

<!-- grade_id: module-integration, points: 15, grade: true, solution: false -->
```python


def test_module():
    """
    Comprehensive test of entire spatial module functionality.

    This final test runs before module summary to ensure:
    - All unit tests pass
    - Functions work together correctly
    - Module is ready for integration with TinyTorch
    """
    print("ðŸ§ª RUNNING MODULE INTEGRATION TEST")
    print("=" * 50)

    # Run all unit tests
    print("Running unit tests...")
    test_unit_conv2d()
    test_unit_pooling()
    test_unit_simple_cnn()

    print("\nRunning integration scenarios...")

    # Test realistic CNN workflow
    print("ðŸ”¬ Integration Test: Complete CNN pipeline...")

    # Create a mini CNN for CIFAR-10
    conv1 = Conv2d(3, 8, kernel_size=3, padding=1)
    pool1 = MaxPool2d(2, stride=2)
    conv2 = Conv2d(8, 16, kernel_size=3, padding=1)
    pool2 = AvgPool2d(2, stride=2)

    # Process batch of images
    batch_images = Tensor(np.random.randn(4, 3, 32, 32))

    # Forward pass through spatial layers
    x = conv1(batch_images)  # (4, 8, 32, 32)
    x = pool1(x)             # (4, 8, 16, 16)
    x = conv2(x)             # (4, 16, 16, 16)
    features = pool2(x)      # (4, 16, 8, 8)

    # Validate shapes at each step
    assert x.shape[0] == 4, f"Batch size should be preserved, got {x.shape[0]}"
    assert features.shape == (4, 16, 8, 8), f"Final features shape incorrect: {features.shape}"

    # Test parameter collection across all layers
    all_params = []
    all_params.extend(conv1.parameters())
    all_params.extend(conv2.parameters())
    # Pooling has no parameters
    assert len(pool1.parameters()) == 0
    assert len(pool2.parameters()) == 0

    # Verify we have the right number of parameter tensors
    assert len(all_params) == 4, f"Expected 4 parameter tensors (2 conv Ã— 2 each), got {len(all_params)}"

    print("âœ… Complete CNN pipeline works!")

    # Test memory efficiency comparison
    print("ðŸ”¬ Integration Test: Memory efficiency analysis...")

    # Compare different pooling strategies (reduced size for faster execution)
    input_data = Tensor(np.random.randn(1, 16, 32, 32))

    # No pooling: maintain spatial size
    conv_only = Conv2d(16, 32, kernel_size=3, padding=1)
    no_pool_out = conv_only(input_data)
    no_pool_size = np.prod(no_pool_out.shape) * 4  # float32 bytes

    # With pooling: reduce spatial size
    conv_with_pool = Conv2d(16, 32, kernel_size=3, padding=1)
    pool = MaxPool2d(2, stride=2)
    pool_out = pool(conv_with_pool(input_data))
    pool_size = np.prod(pool_out.shape) * 4  # float32 bytes

    memory_reduction = no_pool_size / pool_size
    assert memory_reduction == 4.0, f"2Ã—2 pooling should give 4Ã— memory reduction, got {memory_reduction:.1f}Ã—"

    print(f"  Memory reduction with pooling: {memory_reduction:.1f}Ã—")
    print("âœ… Memory efficiency analysis complete!")

    print("\n" + "=" * 50)
    print("ðŸŽ‰ ALL TESTS PASSED! Module ready for export.")
    print("Run: tito module complete 09")
```

<!-- grade_id: main-execution, points: 0, grade: false, solution: true -->
```python
# Run comprehensive module test
if __name__ == "__main__":
    test_module()
```

## ðŸŽ¯ MODULE SUMMARY: Spatial Operations

Congratulations! You've built the spatial processing foundation that powers computer vision!

### Key Accomplishments
- Built Conv2d with explicit loops showing O(NÂ²MÂ²KÂ²) complexity âœ…
- Implemented MaxPool2d and AvgPool2d for spatial dimension reduction âœ…
- Created SimpleCNN demonstrating spatial operation integration âœ…
- Analyzed computational complexity and memory trade-offs in spatial processing âœ…
- All tests pass including complete CNN pipeline validation âœ…

### Systems Insights Discovered
- **Convolution Complexity**: Quadratic scaling with spatial size, kernel size significantly impacts cost
- **Memory Patterns**: Pooling provides 4Ã— memory reduction while preserving important features
- **Architecture Design**: Strategic spatial reduction enables parameter-efficient feature extraction
- **Cache Performance**: Spatial locality in convolution benefits from optimal memory access patterns

### Ready for Next Steps
Your spatial operations enable building complete CNNs for computer vision tasks!
Export with: `tito module complete 09`

**Next**: Milestone 03 will combine your spatial operations with training pipeline to build a CNN for CIFAR-10!

Your implementation shows why:
- Modern CNNs use small kernels (3Ã—3) instead of large ones (computational efficiency)
- Pooling layers are crucial for managing memory in deep networks (4Ã— reduction per layer)
- Explicit loops reveal the true computational cost hidden by optimized implementations
- Spatial operations unlock computer vision - from MLPs processing vectors to CNNs understanding images!